Antes de empezar a resolver los problemas de este notebook, asegúrate de que todo funciona como es debido. Primero, **reinicia el kernel** (En el menú superior, selecciona Kernel$\rightarrow$Restart) y después **ejecuta todas las celdas** (en el menú, selecciona Cell$\rightarrow$Run All).

Asegúrate también de que rellanas todas las partes que pongan `ESCRIBE AQUI TU CODIGO` or "ESCRIBE AQUI TU RESPUESTA", así como tu nombre y el de tu pareja EN NAME_1 y COLLABORATORS respectivamente (entre las comillas).

```python
NAME_1 = ""
COLLABORATORS = ""
```

---

# Ejercicio 1:

Completa las funciones dadas para que pasen los tests que se muestran debajo:

<!-- grade_id: cell-ce887e66c886e184, points: 0, grade: false, solution: true -->
```python
import copy

def anade_fila(matriz):
    # ESCRIBE AQUI TU CODIGO
    raise NotImplementedError()

def anade_columna(matriz):
    # ESCRIBE AQUI TU CODIGO
    raise NotImplementedError()

```

<!-- grade_id: cell-2fb372da06681186, points: 5, grade: true, solution: false -->
```python
m = [[0, 0], [0, 0]]
n = [[3, 2, 5], [1, 4, 7]]
assert(anade_fila(m) == [[0, 0], [0, 0], [0, 0]])
assert(anade_fila(n) == [[3, 2, 5], [1, 4, 7], [0, 0, 0]])
assert(m == [[0, 0], [0, 0]])
assert(n == [[3, 2, 5], [1, 4, 7]])

o = [[3, 2], [5, 1], [4, 7]]
assert(anade_columna(m) == [[0, 0, 0], [0, 0, 0]])
assert(anade_columna(o) == [[3, 2, 0], [5, 1, 0], [4, 7, 0]])

assert(m == [[0, 0], [0, 0]])
assert(o == [[3, 2], [5, 1], [4, 7]])
```

Output:
```
<ipython-input-2-04e1127e2752> in <module>()
      1 m = [[0, 0], [0, 0]]
      2 n = [[3, 2, 5], [1, 4, 7]]
----> 3 assert(anade_fila(m) == [[0, 0], [0, 0], [0, 0]])
      4 assert(anade_fila(n) == [[3, 2, 5], [1, 4, 7], [0, 0, 0]])
      5 assert(m == [[0, 0], [0, 0]])

<ipython-input-1-24c7e46afcba> in anade_fila(matriz)
      3 def anade_fila(matriz):
      4     # ESCRIBE AQUI TU CODIGO
----> 5     raise NotImplementedError()
      6 
      7 def anade_columna(matriz):

NotImplementedError:
```

# Ejercicio 2:
Escribe una función python que actúe como el método update() explicado en clase, pero que en vez de sobreescribir valores para claves repetidas los combine en una lista.

<!-- grade_id: cell-971c92f6e667e316, points: 0, grade: false, solution: true -->
```python
def une(diccionario_1, diccionario_2):
    # ESCRIBE AQUI TU CODIGO
    raise NotImplementedError()
```

<!-- grade_id: cell-82e29fdbb35527de, points: 5, grade: true, solution: false -->
```python
cursos = { "curso1":{"titulo":"Introducción a Python", 
                         "localidad":"Madrid", 
                         "instructor":"Dr. Acula"},
              "curso2":{"titulo":"Python Avanzado",
                         "localidad":"Salamanca",
                         "instructor":"Abel Benito"},
              "curso3":{"titulo":"Procesamiento de Textos",
                         "localidad":"Zamora",
                         "instructor":"Armando Guerra"}
              }

cursos_2 = { "curso4":{"titulo":"Introducción a Python", 
                         "localidad":"Madrid", 
                         "instructor":"Dr. Acula"},
              "curso1":{"titulo":"Python Avanzado",
                         "localidad":"Salamanca",
                         "instructor":"Abel Benito"}
              }

cursos_union = { "curso1":[{"titulo":"Introducción a Python", 
                         "localidad":"Madrid", 
                         "instructor":"Dr. Acula"},
                           {"titulo":"Python Avanzado",
                         "localidad":"Salamanca",
                         "instructor":"Abel Benito"}],
              "curso2":{"titulo":"Python Avanzado",
                         "localidad":"Salamanca",
                         "instructor":"Abel Benito"},
              "curso3":{"titulo":"Procesamiento de Textos",
                         "localidad":"Zamora",
                         "instructor":"Armando Guerra"},
                "curso4":{"titulo":"Introducción a Python", 
                         "localidad":"Madrid", 
                         "instructor":"Dr. Acula"}
              }

assert(une(cursos, cursos_2) == cursos_union)
```