Antes de empezar a resolver los problemas de este notebook, asegúrate de que todo funciona como es debido. Primero, **reinicia el kernel** (En el menú superior, selecciona Kernel$\rightarrow$Restart) y después **ejecuta todas las celdas** (en el menú, selecciona Cell$\rightarrow$Run All).

Asegúrate también de que rellanas todas las partes que pongan `ESCRIBE AQUI TU CODIGO` or "ESCRIBE AQUI TU RESPUESTA", así como tu nombre y el de tu pareja EN NAME_1 y COLLABORATORS respectivamente (entre las comillas).

```python
NAME_1 = ""
COLLABORATORS = ""
```

---

# Práctica 1: Sopa de números

## Discusión 
En esta práctica vas a aplicar lo que has aprendido en los cuatro primeros talleres de la asignatura y en particular el tratamiento de cadenas y los bucles. Te recomiendo que leas este enunciado varias veces para entender bien lo que se pide. Si a pesar de ello aún tienes dudas, puedes consultarme a través del foro de la asignatura creado para comentar esta práctica o reservar una hora de tutoría. **No respondo dudas referentes a la práctica por correo.**

Dada una matriz cuadrada de orden n, se pueden formar secuencias de elementos contiguos en **tres direcciones**: Horizontal, vertical y diagonal y en **ocho sentidos** distintos, respectivamente:
1. Horizontal
    1. Izquierda-derecha.
    2. Derecha-izquierda.
2. Vertical 
    1. Arriba-abajo.
    2. Abajo-arriba.
3. Diagonal
    1. Izquierda-derecha y Arriba-abajo.
    2. Izquierda-derecha y Abajo-arriba.
    3. Derecha-izquierda y Arriba-abajo.
    4. Derecha-izquierda y Abajo-arriba.
    
Como ejemplo, consideremos la siguiente matriz M de dimensión cinco:

$$ M_{1} = 
 \begin{matrix}
  6 & 2 & 9 & 5 & 9 \\
  2 & 9 & 6 & 7 & 8 \\
  4 & 2 & 8 & 8 & 7 \\
  2 & 2 & 7 & 4 & 2 \\
  2 & 2 & 3 & 2 & 2
 \end{matrix} $$

Podremos formar números de cinco o menos dígitos, de los cuales por definición, unos serán pares y otros impares.
En el caso de la matriz M, los cuatro números de cinco dígitos más grandes de toda la matriz (obviamente son aquellos que empiezan por 9) son los que se encuentran en las siguientes posiciones, resultantes de buscar en todas las direcciones y sentidos mencionados anteriormente:

$$
\begin{align}
n_{1} = a_{0,4}a_{1,4}a_{2,4}a_{3,4}a_{4,4} = 98722\\
n_{2} = a_{0,4}a_{1,3}a_{2,2}a_{3,1}a_{4,0} = 97822\\
n_{3} = a_{0,2}a_{1,2}a_{2,2}a_{3,2}a_{4,2} = 96873\\
n_{4} = a_{0,2}a_{1,2}a_{2,2}a_{3,2}a_{4,2} = 95926
\end{align}
$$

De los números obtenidos, $n_{1}, n_{2}, n_{4}$ son pares mientras que $n_{3}$ es impar, por lo tanto el mayor número impar que puedo formar a partir de posiciones contiguas de la matriz es $n_{3}$ y el mayor número par es $n_{1}$.

## Enunciado
Se pide implementar en código la función encuentra_maximo (se proporciona una plantilla en la celda de respuesta) que recibe un único argumento "sopa" y devuelve **el mayor número entero impar** encontrado en la matriz. El formato empleado se puede ver en el siguiente ejemplo, en el que se codifica la matriz M vista anteriormente y se pasa a dicha función, que devuelve la solución (El ejemplo imprime por pantalla 96873).

```python
sopa ='''62959\n29678\n42887\n22742\n22322'''
print(encuentra_maximo(sopa))

sopa ='''62959\n29678\n42887\n22742\n22322'''
print(encuentra_maximo(sopa))
```

Es importante considerar aquellos casos en los que la longitud del resultado no coincide con la dimensión de la matriz. Veamos otro ejemplo con una matriz $M_{2}$ resultante de modificar la empleada en el ejemplo anterior.

$$ M_{2} = 
 \begin{matrix}
  6 & 2 & 6 & 6 & 8 \\
  2 & 9 & 6 & 7 & 8 \\
  4 & 2 & 8 & 8 & 6 \\
  2 & 2 & 7 & 4 & 2 \\
  2 & 2 & 4 & 2 & 2
 \end{matrix} $$

Al variar algunos dígitos no es posible encontrar un número impar de cinco cifras, así que la solución estaría en los de cuatro. Es fácil imaginar situaciones análogas en las que sería necesario aumentar la búsqueda aún más, a números de tres e incluso dos cifras, que tu código deberá resolver también. 

*NOTA: Para terminar la práctica, explica tu solución en la última usando celda de respuesta de texto*


## Consideraciones y consejos
1. Tu solución funcionara para cualquier matriz M de entrada de dimensión $n\geq2$ tal que $ \forall a_{i,j} \in M, a_{i,j} \geq 0$ 
2. Para aprobar la práctica, es necesario que tu código al menos encuentre una solución considerando filas y columnas en todas las direcciones y sentidos posibles.
3. **Cuidado** con los ceros al principio: $12340 > 04321 $
4. Si no existen números impares en toda la matriz, la función *encuentra_maximo* devolverá -1.
5. Se considerará positivamente la claridad, limpieza y simplicidad del código. Cuanto más concisa, elegante y rápida sea tu solución, mejor.
6. Si copias código de Internet, asegúrate de que lo entiendes, puede que tengas que explicarlo en una defensa si tengo dudas de su autoría. 
7. Creo que sobra decirlo pero **no compartáis código fuera de la pareja.**
8. Te aconsejo que muestres el número de línea para encontrar rápidamente los errores que tengas al programar. Para ello vete a: View -> Toggle Line Numbers".
9. Salva tu trabajo a menudo utilizando el icono del disquete de la esquina superior izquierda.
10. Como pista, en la casilla de respuesta te he dejado las signaturas de algunas funciones que considero que son útiles para resolver el problema. Si no sabes por dónde empezar, piensa en qué casos llamarías a estas funciones. Puedes (y debes) declarar tú otras funciones que consideres oportunas para la correcta resolución de la práctica. **El único requisito es que exista una función llamada encuentra_maximo, que reciba un argumento y devuelva el resultado como se ha explicado**.
11. Las soluciones que añadan **comentarios** al código (Usando el caracter almohadilla #) serán evaluadas positivamente.



## Tests
Si usas la opción del notebook "Validate", tu código se evaluará contra una serie de tests, que te darán una pista de cómo de buena es tu solución. Los tests se ejecutan en orden creciente de complejidad (los primeros tests prueban funcionalidad más básica). Para hacerte una idea, tu código debería pasar los **primeros siete** tests para que apruebes la práctica. 
**IMPORTANTE: Que tu código pase los tests no es sinónimo de que la práctica esté aprobada**.


## Cuidado
1. Asegúrate de que tu código cumple los requisitos de esta práctica y que **no da errores**. Los envíos que den error no serán corregidos y la nota obtenida será un 0.
1. **NO** uses la función de entrada por teclado *input*. Puedes usar *print* para imprimir mensajes de debug durante la realización de la práctica, pero asegúrate de eliminarlos cuando vayas entregarla.
2. Emplea sólo lo explicado hasta el lab 3 de la asignatura.
3. No se evaluarán:
    1. Soluciones que den error.
    2. Soluciones que empleen listas, matrices, objetos o cualquier otra característica del lenguaje no explicada en los cuatro primeros laboratorios prácticos de la asignatura (hasta lab 3).
    3. Soluciones que empleen la sentencia **import** para hacer uso de paquetes estándar o de terceros. 
    4. **El límite de entrega estricto es el 20 de Diciembre de 2017 a las 0:00 horas. No se corregirán prácticas entregadas después.**


<!-- grade_id: cell-bd20e5bfdf98ae5d, points: 0, grade: false, solution: true -->
```python

#Extrae la dimensión de la sopa de entrada. 
def extrae_dimension(soup):
    """Extrae la dimensión de la sopa de números"""
    return soup.count('\n') + 1

#Devuelve el caracter en la posicion [fila, columna] de la sopa, dada la dimension de la misma
def char_at_position(soup, dim, fila, columna):
    return soup[(dim+1)*fila + columna]


def encuentra_maximo(sopa):
    return -1
```

<!-- grade_id: cell-d4ed2c868ea8d714, points: 10, grade: true, solution: false -->
```python
#Aprobado
##Sin números impares
assert(encuentra_maximo('''62\n28''') == -1)
##2x2
assert(encuentra_maximo('''62\n27''') == 67)

#5x5
##Columna arriba-abajo
assert(encuentra_maximo('''62959\n29678\n42887\n22742\n22322''') == 96873)
##Columna abajo-arriba
assert(encuentra_maximo('''62656\n29678\n42886\n22742\n22222''') == 24875)
##Fila izquierda-derecha
assert(encuentra_maximo('''62656\n29679\n42886\n22742\n22222''') == 29679)
##Fila derecha-izquierda
assert(encuentra_maximo('''62656\n79679\n42886\n22742\n22222''') == 97697)

#10x10

assert(encuentra_maximo('''6743605871\n1820953321\n8958054851\n7486519455\n1389386144\n4672017818\n9082874626\n6728050071\n1684990063\n4374367650''') == 8958054851)


##Diagonales
assert(encuentra_maximo('''92656\n79679\n42886\n22782\n22221''') == 99881)
assert(encuentra_maximo('''32651\n71674\n42885\n22742\n72223''') == 72871)
assert(encuentra_maximo('''32659\n79674\n42886\n22742\n12229''') == 97821)
assert(encuentra_maximo('''32639\n79674\n42886\n22742\n22229''') == 94893)

#Soluciones con m < n
assert(encuentra_maximo('''41622\n24084\n42084\n27746\n40220''') == 7241)
assert(encuentra_maximo('''42622\n24684\n42084\n22742\n42220''') == 6607)
assert(encuentra_maximo('''44444\n44444\n44944\n44444\n44444''') == 449)
assert(encuentra_maximo('''4444\n4444\n4944\n4444''') == 449)

```

<!-- grade_id: cell-d150ccb27c7334e8, points: 0, grade: true, solution: true -->
ESCRIBE AQUI TU RESPUESTA