Antes de empezar a resolver los problemas de este notebook, asegúrate de que todo funciona como es debido. Primero, **reinicia el kernel** (En el menú superior, selecciona Kernel$\rightarrow$Restart) y después **ejecuta todas las celdas** (en el menú, selecciona Cell$\rightarrow$Run All).

Asegúrate también de que rellanas todas las partes que pongan `ESCRIBE AQUI TU CODIGO` or "ESCRIBE AQUI TU RESPUESTA", así como tu nombre y el de tu pareja EN NAME_1 y COLLABORATORS respectivamente (entre las comillas).

```python
NAME_1 = ""
COLLABORATORS = ""
```

---

# Ejercicio 1:
Escribe una función que devuelva la suma de todos los elementos de una lista.

<!-- grade_id: cell-d02b0c4baeb11dae, points: 0, grade: false, solution: true -->
```python
def suma(lista):
    acc = 0
    for e in lista:
        acc += e
    return acc
```

<!-- grade_id: cell-5272a50bddf5bfbf, points: 5, grade: true, solution: false -->
```python
assert(suma([1,2,3,4,5]) == 15)
assert(suma([1,2,3,4,5,6,7]) == 28)

```

# Ejercicio 2:
Escribe la función en python que recibe una tupla de elementos y devuelve otra tupla que contiene los mismos elementos en grupos del número que se proporciona por parámetros

<!-- grade_id: cell-05e1ea1ae39328a9, points: 0, grade: false, solution: true -->
```python
def agrupa(tupla, numero):
    res = [] 
    for i in range(0, len(tupla), numero):
        interm = []
        for j in range(numero):
            if i+j < len(tupla):
                interm.append(tupla[i+j])
        res.append(tuple(interm))
    print(tuple(res))
    return tuple(res)
```

```python
def agrupa(tupla, numero):
    nueva_lista = tuple()
    temp = tuple()
    
    #Recorremos los números de la lista
    for i in range(0, len( tupla)):
        #los añadimos a una lista mas pequeña
        temp = temp + (tupla[i],)
        
        #cuando hay "número"(el número por el que esté formado cada pequeña lista), añadimos la lista menor a una lista de listas(la lista mayor)
        if((i+1)%numero == 0):
            nueva_lista += ((temp),)
            #a continuación vaciamos la lista pequeña
            temp = tuple()
    #los numeros sueltos que queden pq no llegan a llenar un conjunto entero, los añadiremos tal cual (el 10 al hacer grupos de 3)        
    if len(tupla)%numero != 0:
        nueva_lista = nueva_lista + tuple(temp)
    #print no es necesario al igual que en el caso anterior
    print(nueva_lista)
    return(nueva_lista)
```

<!-- grade_id: cell-929ffaa649a30419, points: 5, grade: true, solution: false -->
```python
assert(agrupa((1,2,3,4,5,6,7,8,9,10), 2) == ((1,2),(3,4),(5,6),(7,8),(9,10)))
assert(agrupa((1,2,3,4,5,6,7,8,9,10), 3) == ((1,2,3),(4,5,6),(7,8,9),(10)))
```

Output:
```
((1, 2), (3, 4), (5, 6), (7, 8), (9, 10))
((1, 2, 3), (4, 5, 6), (7, 8, 9), 10)
```

```python
assert(agrupa((1,2,3,4,5,6,7,8,9,10), 3) == ((1,2,3),(4,5,6),(7,8,9),(10,)))
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-7-7ed071fa50ae> in <module>()
----> 1 assert(agrupa((1,2,3,4,5,6,7,8,9,10), 3) == ((1,2,3),(4,5,6),(7,8,9),(10,)))

AssertionError:
```

```python

```