Antes de empezar a resolver los problemas de este notebook, asegúrate de que todo funciona como es debido. Primero, **reinicia el kernel** (En el menú superior, selecciona Kernel$\rightarrow$Restart) y después **ejecuta todas las celdas** (en el menú, selecciona Cell$\rightarrow$Run All).

Asegúrate también de que rellanas todas las partes que pongan `ESCRIBE AQUI TU CODIGO` or "ESCRIBE AQUI TU RESPUESTA", así como tu nombre y el de tu pareja EN NAME_1 y COLLABORATORS respectivamente (entre las comillas).

```python
NAME_1 = ""
COLLABORATORS = ""
```

---

## Ejercicio 1: PimPam
Escribe el código de un bucle que imprima por pantalla los números del 1 al 20. Para aquellos números que sean múltiplos de tres, imprime "Pim" (Sin las comillas) en vez del número. Para los múltiplos de cinco, imprime "Pam". Para aquellos números que sean a la vez múltiplos de tres y de cinco imprime "PimPam". 

<!-- grade_id: cell-8a9d41e30e850c8e, points: 5, grade: true, solution: true -->
```python
# ESCRIBE AQUI TU CODIGO
raise NotImplementedError()
```

## Ejercicio 2: Replace
Con lo que hemos dado en clase hoy, ¿podrías tú implementar la función str.replace?
Esto es, tienes que implementar una función que reciba tres argumentos: una cadena y dos caracteres. La función ha de devolver una nueva cadena en la que todas las apariciones del primer caracter hayan sido reemplazadas por el segundo.

<!-- grade_id: cell-b73077926760d16d, points: 0, grade: false, solution: true -->
```python
def my_replace(cadena, caracter1, caracter2):
    # ESCRIBE AQUI TU CODIGO
    raise NotImplementedError()
```

<!-- grade_id: cell-3ed3f8a08b4aaf65, points: 0, grade: true, solution: false -->
```python
assert my_replace('cadena', 'a', 'o') == 'codeno'
assert my_replace('esto es una prueba', 'a', 'o') == 'esto es uno pruebo'
```

```python

```