Antes de empezar a resolver los problemas de este notebook, aseg칰rate de que todo funciona como es debido. Primero, **reinicia el kernel** (En el men칰 superior, selecciona Kernel$\rightarrow$Restart) y despu칠s **run all cells** (en el men칰, selecciona Cell$\rightarrow$Run All).

Aseg칰rate tambi칠n de que rellanas todas las partes que pongan `ESCRIBE AQUI TU CODIGO` or "ESCRIBE AQUI TU RESPUESTA", as칤 como tu nombre y el de tu pareja EN NAME_1 y COLLABORATORS respectivamente (entre las comillas).

```python
NAME_1 = ""
COLLABORATORS = ""
```

---

![](jupyter.png)

Con este notebook vamos a explicar el funcionamiento de los notebooks entregables sobre los que se realizar치 el seguimiento de las pr치cticas semanales. En ellos se plantear치n problemas sobre los conceptos explicados en clase, y se han de entregar durante la segunda hora de pr치cticas. 

La particularidad de estos notebooks es que contienen **"celdas de tarea o pregunta"**. Estas celdas est치n dise침adas para que practiques lo aprendido en la primera hora de clase. Para ello, cada una de estas celdas almacenar치 las respuestas a los problemas que tendr치s que enviarme al final de cada sesi칩n. 

Cada notebook entregable tiene un total de puntos m치ximos a conseguir que variar치 dependiendo de la sesi칩n. Este notebook por ejemplo tiene **6 puntos**.

Vamos a explicar ahora los distintos tipos de celdas que te podr치s encontrar en este tipo de cuestionarios.

---
## Celda de respuesta de texto (2 puntos)

En estas celdas se espera que escrib치is una respuesta en texto a una pregunta dada. La respuesta ser치 evaluada por el profesor hasta un m치ximo de los puntos indicados (en este caso 2). 
Para un primer ejemplo sencillo, reemplazad el texto en la siguiente celda por vuestro nombres completos:

<!-- grade_id: cell-b4c5d2801ddd2f91, points: 2, grade: true, solution: true -->
ESCRIBE AQUI TU RESPUESTA

## Celda de respuesta de c칩digo (2 puntos)

En este segundo tipo de celdas se espera que escrib치is c칩digo (a ser posible que funcione!). Este c칩digo a su vez, tendr치 que realizar ciertas funcionames y normalmente se comprobar치 en las siguientes celdas que as칤 sea.
En este tipo de celdas podr치s comprobar que tu respuesta es v치lida. No te recomiendo entregar ejercicios cuya respuesta sea v치lida: Se evaluar치n negativamente. Si te atascas consulta a tus compa침eros o p칤deme ayuda.

Para ilustrar este caso, toma el c칩digo del algoritmo de Euclides en Python visto en clase de teor칤a.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def mcd(a, b):
    """Algoritmo de Euclides para el c치lculo del M.C.D
    """
    # ESCRIBE AQUI TU CODIGO
    raise NotImplementedError()
```

Tu c칩digo deber칤a sacar por pantalla `4` para a=12 y b=8. Comprueba que es as칤:

```python
mcd(2,0)
```

<!-- grade_id: correct_squares, points: 2, grade: true, solution: false -->
```python
"""Comprueba que la salida es correcta para un cierto n칰mero de entradas distintas"""
assert mcd(8, 12) == 4
assert mcd(2, 2) == 2
assert mcd(0, 2) == 2
assert mcd(0, 0) == 0
assert mcd(100, 1000) == 100
```

---

## Rizando el rizo 游 (3 puntos)
Para terminar, normalmente os pedir칠 que complet칠is un ejercicio de los de "idea feliz", para asegurarme de que os v치is de clase con todo aprendido.

Por ejemplo, si te digo que el mcm se define como: 
\begin{align}
mcd(a,b) & = \frac{a \cdot b}{MCM(a,b)} \\
\end{align}


쯇odr칤as cambiar la funci칩n *mcd* para que devuelva el m칤nimo com칰n m칰ltiplo de a y b?

<!-- grade_id: sum_of_squares_application, points: 0, grade: false, solution: true -->
```python
def mcm(a, b):
    """Funci칩n para el c치lculo del M.C.M. de 2 n칰mero (basada en el M.C.D.)
        """
    num = a * b
    # ESCRIBE AQUI TU CODIGO
    raise NotImplementedError()
```

```python
mcm(2, 4)
```

<!-- grade_id: cell-e3279643d328ccd7, points: 2, grade: true, solution: false -->
```python
assert mcm(8, 12) == 24
assert mcm(2, 2) == 2
assert mcm(0, 2) == 0
assert mcm(0, 0) == 0
assert mcm(100, 1000) == 1000
```

```python

```